---
id: 5f33071498eb2472b87ddee4
title: Part 1
challengeType: 0
dashedName: part-1
---

# --description--

ตามที่คุณได้เรียนไปแล้วในขั้นตอนที่ผ่านมาของ Cat Photo App มันมีโครงสร้างพื้นฐานที่จำเป็นในการสร้างหน้าเว็บเพจ

จงเพิ่ม `<!DOCTYPE html>` tag, และ `html` element อันหนึ่ง

# --hints--

คุณควรมี `DOCTYPE` declaration.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

You should have an opening `<html>` tag.

```js
assert(code.match(/<html>/i));
```

You should have a closing `<html>` tag. Remember that closing tags have a `/` following the opening `<` bracket.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
