---
id: 5d8a4cfbe6b6180ed9a1c9de
title: Part 1
challengeType: 0
dashedName: part-1
---

# --description--

สำหรับ dashboard project! จะใช้ JavaScript data visualization library, D3 เพื่อสร้าง visualization สำหรับ social media ที่ประกอบไปด้วย line graph, a pie chart, และ legend

ขั้นแรกให้สร้าง HTML file โดยเริ่มจากการเขียน `<!DOCTYPE html>` ตรงส่วนหัวของ file เพื่อบอก browser ให้รู้จักกับประเภทของ document

# --hints--

test-text

```js
assert(/<!DOCTYPE\s+html\s*>/gi.test(code));
```

# --seed--

## --seed-contents--

```html
```

# --solutions--

```html
<!DOCTYPE html>
```
