---
id: 587d8256367417b2b2512c78
title: Adjacency Matrix
challengeType: 1
forumTopicId: 301621
dashedName: adjacency-matrix
---

# --description--

อีกวิธีนึงในการเก็บข้อมูลของกราฟคือการใช้ <dfn>adjacency matrix</dfn>  
ซึ่ง <dfn>adjacency matrix</dfn> นี้คือ array สองมิติ (2D) โดยที่แต่ละ array ชั้นในมีจำนวน element เท่ากันกับ array ชั้นนอก 
ซึ่งก็คือเมทริกซ์หรือตารางของตัวเลข โดยที่ตัวเลขเป็นตัวแทนของความสัมพันธ์

**หมายเหตุ**: ตัวเลขด้านบนและด้านซ้ายของเมทริกซ์เป็นข้อมูลที่แสดงถึง node  

โดยที่ภายในเมทริกซ์ 1 หมายถึงมีความสัมพันธ์ระหว่างแถวและคอลัมน์ และ 0 หมายถึงความสัมพันธ์

<pre>
    1 2 3
  \------
1 | 0 1 1
2 | 1 0 0
3 | 1 0 0
</pre>

ด้านบนเป็นกราฟธรรมดาๆที่ไม่มีทิศทาง ซึ่งมีสาม node โดยที่ node แรกเชื่อมต่อกับ node ที่สองและสาม  
ส่วนด้านล่างเป็นการแปลภาพด้านบนให้เป็น JavaScript:

```js
var adjMat = [
  [0, 1, 1],
  [1, 0, 0],
  [1, 0, 0]
];
```

เมตริกซ์จะไม่เหมือนกับ adjacency list เพราะว่าแต่ละ "แถว" ของเมทริกซ์ต้องมี element จำนวนเท่ากันกับ node ในกราฟ 

ในตัวอย่างนี้ เรามีเมทริกซ์ขนาด 3x3 แปลว่าในกราฟของเราจะมีสาม node โดยกราฟที่มีทิศทางจะมีลักษณะคล้ายกัน 

ด้านล่างเป็นกราฟทื่มีทิศทาง node แรกมีความสัมพันธ์โดยชี้ไปยัง node ที่สอง และ node ที่สองมีความสัมพันธ์โดยชี้ไปยัง node ที่สาม

```js
var adjMatDirected = [
  [0, 1, 0],
  [0, 0, 1],
  [0, 0, 0]
];
```

กราฟสามารถมี <dfn>weight (น้ำหนัก)</dfn> ที่ความสัมพันธ์ได้ด้วย  
ตอนนี้กราฟของเรามีความสัมพันธ์แบบ <dfn>unweighted (ไม่มีน้ำหนัก)</dfn> ซึ่งจะเก็บแค่ว่ามีความสัมพันธ์หรือไม่มีความสัมพันธ์อยู่เท่านั้น โดยแสดงความสัมพันธ์โดยใช้เลขฐานสอง (`0` หรือ `1`)  
ซึ่งคุณสามารถกำหนดน้ำหนักของความสัมพันธ์ได้ ตามการใช้งานของคุณได้เลย (เช่นให้ความสัมพันธ์มีค่าเป็น `2` `3` ฯลฯ)

# --instructions--

ให้สร้าง adjacency matrix ของกราฟแบบไม่มีทิศทางที่มีห้า node 

เมทริกซ์นี้ต้องเป็น array หลายมิติ โดนที่ node ทั้งห้านี้มีความสัมพันธ์ระหว่าง node ดังนี้:  
node ที่หนึ่งสัมพันธ์กับ node ที่สี่  
node ที่หนึ่งสัมพันธ์กับ node ที่สาม  
node ที่สามสัมพันธ์กับ node ที่ห้า  
node ที่สี่สัมพันธ์กับ node ที่ห้า 

โดยที่น้ำหนักของความสัมพันธ์ของทุก node เป็น 1

# --hints--

`undirectedAdjList` ต้องมีทั้งหมด 5 node

```js
assert(
  adjMatUndirected.length === 5 &&
    adjMatUndirected
      .map(function (x) {
        return x.length === 5;
      })
      .reduce(function (a, b) {
        return a && b;
      })
);
```

ต้องมีความสัมพันธ์ระหว่าง node แรกกับ node ที่สี่

```js
assert(adjMatUndirected[0][3] === 1 && adjMatUndirected[3][0] === 1);
```

ต้องมีความสัมพันธ์ระหว่าง node แรกกับ node ที่สาม

```js
assert(adjMatUndirected[0][2] === 1 && adjMatUndirected[2][0] === 1);
```

ต้องมีความสัมพันธ์ระหว่าง node ที่สามกับ node ที่ห้า

```js
assert(adjMatUndirected[2][4] === 1 && adjMatUndirected[4][2] === 1);
```

ต้องมีความสัมพันธ์ระหว่าง node ที่สี่กับ node ที่ห้า

```js
assert(adjMatUndirected[3][4] === 1 && adjMatUndirected[4][3] === 1);
```

# --seed--

## --seed-contents--

```js
var adjMatUndirected = [];
```

# --solutions--

```js
var adjMatUndirected = [
  [0, 0, 1, 1, 0],
  [0, 0, 0, 0, 0],
  [1, 0, 0, 0, 1],
  [1, 0, 0, 0, 1],
  [0, 0, 1, 1, 0]
];
```
