---
id: 587d8256367417b2b2512c79
title: Incidence Matrix
challengeType: 1
forumTopicId: 301644
dashedName: incidence-matrix
---

# --description--

อีกวิธีหนึ่งในการเก็บข้อมูลของกราฟการใช้ <dfn>incidence matrix</dfn>

<dfn>incidence matrix</dfn> คือ array สองมิติ (2D)  
ซึ่ง incidence matrix จะเก็บค่าความสัมพันธ์ของคลาสของ object ที่แตกต่างกันในสองมิติ  
เมทริกซ์ประเภทนี้จะคล้ายกับ adjacency matrix แต่ความหมายของแถวและคอลัมน์จะแตกต่างออกไป

ในกราฟ เราจะมีความสัมพันธ์และ node โดยความสัมพันธ์และ node นี้จะเป็น "คลาสของ object ที่แตกต่างกัน" ของเรา เมทริกซ์นี้จะมีแถวเป็น node และคอลัมน์เป็น edge  
แปลว่าเราสามารถมีจำนวนแถวและคอลัมน์ไม่เท่ากันได้

แต่ละคอลัมน์จะแสดงความสัมพันธ์ที่ไม่ซ้ำกัน และแต่ละความสัมพันธ์ยังเชื่อมต่อสอง node เข้าด้วยกันด้วย 

ในการแสดงความสัมพันธ์ระหว่างสอง node ตัว คุณจะต้องใส่ 1 ในทั้งสองแถวของคอลัมน์ใดคอลัมน์หนึ่ง 

ตัวอย่างด้านล่างเป็นกราฟที่มีสาม node ซึ่งมีความสัมพันธ์เดียว คือความสัมพันธ์ระหว่าง node 1 และ node 3

<blockquote>    1<br>   ---<br>1 | 1<br>2 | 0<br>3 | 1</blockquote>

ด้านล่างเป็นตัวอย่างของ `incidence matrix` ที่มี 4 ความสัมพันธ์และมี 4 node (อย่าลืมว่า คอลัมน์คือความสัมพันธ์ และแถวคือ node):

<blockquote>    1 2 3 4<br>   --------<br>1 | 0 1 1 1<br>2 | 1 1 0 0<br>3 | 1 0 0 1<br>4 | 0 0 1 0</blockquote>

ในตัวอย่างด้านบนจะทำให้เกิดความสัมพันธ์ระหว่าง node ดังนี้:  
2 กับ 3  
1 กับ 2  
1 กับ 4  
1 กับ 3

ซึ่งถ้านำตัวอย่างด้านบนมาเขียนเป็น JavaScript จะได้เป็น:

```js
var incMat = [
  [0, 1, 1, 1],
  [1, 1, 0, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
];
```

ในการสร้างกราฟที่มีทิศทาง เราจะใช้ `-1` เพื่อแสดงความสัมพันธ์ที่มีทิศทางออกจาก node และ `1` แทนความสัมพัน์ที่เข้าสู่ node 

```js
var incMatDirected = [
  [ 0, -1,  1, -1],
  [-1,  1,  0,  0],
  [ 1,  0,  0,  1],
  [ 0,  0, -1,  0]
];
```

ความสัมพันธ์ด้านบนสามารถอธิบายได้ง่ายๆดังนี้:  
2 -> 3  
1 -> 2  
4 -> 1  
1 -> 3

และเราก็สามารถให้ <dfn>weight (น้ำหนัก)</dfn> กับความสัมพันธ์ได้ด้วย ซึ่งในตัวอย่างด้านบนเป็นตัวอย่างที่ <dfn>unweighted (ไม่มีน้ำหนัก)</dfn> ในความสัมพันธ์ ที่ทำเพียงแค่การแสดงความสัมพันธ์โดยใช้เลขฐานสอง (`0` หรือ `1`)  
ซึ่งคุณสามารถกำหนดน้ำหนักของความสัมพันธ์ได้ ตามการใช้งานของคุณได้เลย (เช่นให้ความสัมพันธ์มีค่าเป็น `2` `3` ฯลฯ)

# --instructions--

ให้สร้าง incidence matrix ของกราฟแบบไม่มีทิศทาง โดยให้มีห้า node และสี่ความสัมพันธ์ ซึ่งเมทริกซ์นี้ต้องเป็น array หลายมิติ

node ทั้งห้านี้ต้องมีความสัมพันธ์ดังต่อไปนี้:  

- ความสัมพันธ์แรกเป็นความสัมพันธ์ระหว่าง node แรกและ node ที่สอง  
- ความสัมพันธ์ที่สองเป็นความสัมพันธ์ระหว่าง node ที่สองและ node ที่สาม  
- ความสัมพันธ์ที่สามเป็นความสัมพันธ์ระหว่าง node ที่สามและ node ที่ห้า  
- ความสัมพันธ์ที่สี่เป็นความสัมพันธ์ระหว่าง node ที่สี่และ node ที่สอง  

โดยที่น้ำหนักของความสัมพันธ์ทั้งหมดมีค่าเป็น 1  
และลำดับของความสัมพันธ์ก็มีความสำคัญ

# --hints--

`incMatUndirected` ต้องมีแค่ 5 node เท่านั้น

```js
assert(
  incMatUndirected.length === 5 &&
    incMatUndirected
      .map(function (x) {
        return x.length === 4;
      })
      .reduce(function (a, b) {
        return a && b;
      })
);
```

ความสัมพันธ์แรกต้องเป็นความสัมพันธ์ระหว่าง node แรกและ node ที่สอง

```js
assert(incMatUndirected[0][0] === 1 && incMatUndirected[1][0] === 1);
```

ความสัมพันธ์ที่สองต้องเป็นความสัมพันธ์ระหว่าง node ที่สองและ node ที่สาม

```js
assert(incMatUndirected[1][1] === 1 && incMatUndirected[2][1] === 1);
```

ความสัมพันธ์ที่สามต้องเป็นความสัมพันธ์ระหว่าง node ที่สามและ node ที่ห้า

```js
assert(incMatUndirected[2][2] === 1 && incMatUndirected[4][2] === 1);
```

ความสัมพันธ์ที่สี่ต้องเป็นความสัมพันธ์ระหว่าง node ที่สี่และ node ที่สอง

```js
assert(incMatUndirected[1][3] === 1 && incMatUndirected[3][3] === 1);
```

# --seed--

## --seed-contents--

```js
var incMatUndirected = [

];
```

# --solutions--

```js
var incMatUndirected = [[1, 0, 0, 0],[1, 1, 0, 1],[0, 1, 1, 0],[0, 0, 0, 1],[0, 0, 1, 0]];
```
