---
id: 587d7fb1367417b2b2512bf2
title: Use the .env File
challengeType: 2
forumTopicId: 301521
dashedName: use-the--env-file
---

# --description--

ไฟล์ `.env` เป็นไฟล์ลับ ที่ใช้ในการส่งผ่านตัวแปรสภาพแวดล้อม(environment variable)ไปยังแอปพลิเคชันของคุณ ซึ่งไฟล์ลับนี้ ไม่มีใครสามารถเข้าถึงได้นอกจากตัวคุณ เองและมันยังสามารถใช้ในการจัดเก็บข้อมูลส่วนตัวหรือซ่อนไว้ของคุณได้อีกด้วย ตัวอย่างเช่น คุณสามารถจัดเก็บคีย์ API จากการทำงานภายนอกหรือ URI ของฐานข้อมูลของคุณ คุณยังสามารถใช้เพื่อจัดเก็บองค์ประกอบตัวเลือกได้อีกด้วย ด้วยการตั้งค่าองค์ประกอบตัวเลือก คุณสามารถเปลี่ยนลักษณะการทำงานของแอปพลิเคชันได้ โดยไม่ต้องเขียนโค้ดใหม่


ตัวแปรสภาพแวดล้อมสามารถเข้าถึงได้จากแอปในชื่อ `process.env.VAR_NAME` 
ออบเจ็กต์ `process.env` เป็นอ็อบเจ็กต์ของ NODE และเป็นตัวแปรจะถูกส่งผ่านเป็นสตริง โดยปกติแล้ว ชื่อตัวแปรจะเป็นตัวพิมพ์ใหญ่ทั้งหมดและคั่นด้วยอันเดอร์สกร์อ(_) 
ไฟล์ `.env` เป็นไฟล์เชลล์ คุณจึงไม่จำเป็นต้องใส่ชื่อหรือเครื่องหมายคำพูด แต่สิ่งสำคัญที่ควรทราบคือต้องไม่มีที่ว่างรอบเครื่องหมายเท่ากับเมื่อคุณกำหนดค่าให้กับตัวแปรของคุณ เช่น `VAR_NAME=value` โดยปกติ คุณจะใส่นิยามของตัวแปรแต่ละตัวในบรรทัดที่แยกจากกัน

# --instructions--

เรามาทำการใส่ตัวแปรสภาพแวดล้อมลงในองค์ประกอบตัวเลือกกัน!

ลองสร้างไฟล์ `.env` ตรงรูธของโปรเจคไดเรกทอรี่ดู แล้วจัดเก็บตัวแปรไว้ใน `MESSAGE_STYLE=uppercase`

จากนั้น ในตัวจัดการเส้นทาง `/json` ของ GET ที่คุณสร้างในโจทย์ท้าทายล่าสุด แปลงข้อความของอ็อบเจ็กต์ตอบกลับเป็นตัวพิมพ์ใหญ่ ถ้า `process.env.MESSAGE_STYLE` เท่ากับ `uppercase` ถ้าเป็นอย่างนั้น ออบเจ็กต์การตอบกลับควรเป็น `{"message": "Hello json"}` หรือ `{"message": "HELLO JSON"}` ทั้งนี้ขึ้นอยู่กับค่า `MESSAGE_STYLE`

**Note:** ถ้าคุณใช้ Replit คุณจะไม่สามารถสร้างไฟล์ `.env` ได้ ให้สร้างtab <dfn>SECRETS</dfn> เพื่อเพิ่มตัวแปรแทน

ถ้าคุณกำลังทำงานเฉพาะที่ คุณจำเป็นต้องใช้ `dotenv` มันจะช่วยโหลดตัวแปรสภาพแวดล้อมจากไฟล์ `.env` ไปยัง `process.env` ในการติดตั้งทำได้โดยใช้ `npm install dotenv` หลังจากนั้นใส่ `myApp.js` ที่บนหัวของไฟล์ นำเข้าและโหลดตัวแปรพร้อมกับคำสั่ง `require('dotenv').config()`

# --hints--

การตอบสนองจุดปลายทาง `/json` ควรเปลี่ยนตามตัวแปรสภาพแวดล้อม `MESSAGE_STYLE`

```js
(getUserInput) =>
  $.get(getUserInput('url') + '/_api/use-env-vars').then(
    (data) => {
      assert.isTrue(
        data.passed,
        'The response of "/json" does not change according to MESSAGE_STYLE'
      );
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

# --solutions--

```js
/**
  Backend challenges don't need solutions, 
  because they would need to be tested against a full working project. 
  Please check our contributing guidelines to learn more.
*/
```
