---
id: 587d78aa367417b2b2512aed
title: Declare the Doctype of an HTML Document
challengeType: 0
videoUrl: "https://scrimba.com/p/pVMPUv/cra98AJ"
forumTopicId: 301095
dashedName: declare-the-doctype-of-an-html-document
---

# --description--

ในแบบฝึกหัดที่ผ่านๆมา เราได้สอนเรื่อง HTML element ที่แต่ละตัวมีหน้าที่เฉพาะของตัวเอง และได้สอนเรื่องวิธีใช้ element เหล่านั้นด้วย
แต่ก็ยังมี element ที่เราไม่ได้เรียน ซึ่งใช้กำหนดโครงสร้างพื้นฐานของเว็บไซต์ของคุณ โดยใน HTML ทุก ๆ ไฟล์ควรมี element นี้

ที่ด้านบนของไฟล์ คุณจำเป็นต้องบอกให้เบราว์เซอร์รู้ว่า HTML ของหน้าเว็บคุณเป็นเวอร์ชั่นไหน
HTML เป็นภาษาที่อัปเดทค่อนข้างบ่อย และเบราว์เซอร์หลัก ๆ ส่วนมากจะ support HTML5 ซึ่งเป็นเวอร์ชั่นล่าสุดอยู่แล้ว
แต่ก็ยังมีเว็บเก่า ๆ ที่ใช้ HTML เวอร์ชั่นเก่าอยู่บ้าง

วิธีการบอกเบราว์เซอร์ว่าคุณใช้ HTML เวอร์ชันไหนสามารถทำได้โดยใช้เแท็ก `<!DOCTYPE ...>` เขียนไว้ที่บรรทัดแรกของไฟล์ HTML
โดยส่วนที่เป็น `...` จะใช้เพื่อบอกเวอร์ชั่นของ HTML ที่ใช้อยู่
ถ้าเป็น HTML5 เราจะใช้เป็น `<!DOCTYPE html>`

เครื่องหมาย `!` และ `DOCTYPE` ที่เป็นตัวพิมพ์ใหญ่นั้นสำคัญมาก โดยเฉพาะเบราว์เซอร์รุ่นเก่า ๆ

ส่วน `html` นั้นจะเป็นตัวเล็กหรือตัวใหญ่ก็ได้ (case insensitive)

ต่อมา โค้ด HTML ทั้งหมดที่คุณเขียนจะต้องอยู่ในแท็ก `html`
โดยแท็กเปิด `<html>` จะอยู่บรรทัดต่อจาก `<!DOCTYPE html>` ทันที และแท็กปิด `</html>` จะเป็นแท็กสุดท้ายของหน้านี้

ลองดูตัวอย่างโครงสร้างหน้าเว็บไซต์

โค้ด HTML จะต้องอยู่ระหว่างแท็ก `html` ทั้ง 2 อัน

```html
<!DOCTYPE html>
<html></html>
```

# --instructions--

ให้เพิ่มแท็ก `DOCTYPE` เพื่อบอกเบราว์เซอร์ว่าเว็บของคุณเขียนด้วย HTML5 ที่ด้านบนสุดของไฟล์ HTML ใน code editor
และใต้แท็กนี้ ให้แพิ่มแท็ก `html` โดยในแท็ก `html` ตัวนี้จะต้องมี `h1` element อยู่หนึ่งตัว
โดยใน `h1` นั้นคุณจะใส่ข้อความอะไรลงไปก็ได้

# --hints--

โค้ดของคุณต้องมีแท็ก `<!DOCTYPE html>`

```js
assert(code.match(/<!DOCTYPE\s+?html\s*?>/gi));
```

ต้องมี `html` element หนึ่งตัว

```js
assert($("html").length == 1);
```

ในแท็ก `html` ต้องมี `h1` element หนึ่งตัว

```js
assert(code.match(/<html>\s*?<h1>\s*?.*?\s*?<\/h1>\s*?<\/html>/gi));
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <h1>Hello world</h1>
</html>
```
