---
id: 5d822fd413a79914d39e98c9
title: Part 1
challengeType: 0
dashedName: part-1
---

# --description--

ยินดีต้อนรับเข้าสู่โปรเจค CSS Variables Skyline! 
จงเริ่มด้วยการเพิ่ม `!DOCTYPE html` declaration ที่ด้านบนของเอกสารเพื่อที่ browser จะได้รู้ว่ามันกำลังอ่านเอกสารประเภทไหนอยู่

# --hints--

โค้ดของคุณควรมี `DOCTYPE` reference.

```js
assert(code.match(/<!DOCTYPE/gi));
```

You should include a space after the `DOCTYPE` reference.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

You should define the document type to be `html`.

```js
assert(code.match(/html/gi));
```

You should close the `DOCTYPE` declaration with a `>` after the type.

```js
assert(code.match(/html\s*>/gi));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
