---
id: a8d97bd4c764e91f9d2bda01
title: Binary Agents
challengeType: 5
forumTopicId: 14273
dashedName: binary-agents
---

# --description--

จงทำการแปลง binary string ให้กลายเป็นประโยคภาษาอังกฤษ

binary string จะมีเว้นวรรคคั่น

# --hints--

`binaryAgent("01000001 01110010 01100101 01101110 00100111 01110100 00100000 01100010 01101111 01101110 01100110 01101001 01110010 01100101 01110011 00100000 01100110 01110101 01101110 00100001 00111111")` ต้องคืนค่าเป็น string `Aren't bonfires fun!?`

```js
assert.deepEqual(
  binaryAgent(
    '01000001 01110010 01100101 01101110 00100111 01110100 00100000 01100010 01101111 01101110 01100110 01101001 01110010 01100101 01110011 00100000 01100110 01110101 01101110 00100001 00111111'
  ),
  "Aren't bonfires fun!?"
);
```

`binaryAgent("01001001 00100000 01101100 01101111 01110110 01100101 00100000 01000011 01100001 01110010 01100101 01100101 01110010 01010110 01101001 01101111")` ต้องคืนค่าเป็น string `I love CareerVio!`

```js
assert.deepEqual(
  binaryAgent(
    '01001001 00100000 01101100 01101111 01110110 01100101 00100000 01000011 01100001 01110010 01100101 01100101 01110010 01010110 01101001 01101111'
  ),
  'I love CareerVio!'
);
```

# --seed--

## --seed-contents--

```js
function binaryAgent(str) {
  return str;
}

binaryAgent("01000001 01110010 01100101 01101110 00100111 01110100 00100000 01100010 01101111 01101110 01100110 01101001 01110010 01100101 01110011 00100000 01100110 01110101 01101110 00100001 00111111");
```

# --solutions--

```js
function binaryAgent(str) {
  return str.split(' ').map(function(s) { return parseInt(s, 2); }).map(function(b) { return String.fromCharCode(b);}).join('');
}
```
