---
id: 618b97adeefd4004bcd3286d
title: แนะนำตัวเองใน Webboard
challengeType: 2
forumTopicId: 28
dashedName: part-2
---

# --description--

**จุดมุ่งหมาย:** เนื่องจากเราเป็น Project ที่มีผู้มาเรียนเป็นจำนวนมา จึงควรรู้จักกันไว้

ให้ไปที่ [link นี้](https://forum.careervio.com/t/topic/37) แล้วให้ทุกคนแนะนำตัวเอง โดยการ Reply กระทู้link นั้น มาเลยครับ โดยบอกข้อมูลราวๆ นี้

1) ชื่อจริง ชื่อเล่น (หรือนามแฝง กรณีไม่อยากเปิดเผยตัวตน)
2) จุดมุ่งหมายที่มาเรียน
3) อื่นๆ ที่อยากแบ่งปัน

จากนั้น พิมพ์คำว่า "Post แล้ว" ใส่ในช่องคำตอบ 

# --hints--

หลังจาก Post ที่ Webboard เสร็จแล้ว ให้พิมพ์คำว่า "Post แล้ว" ใส่ในช่องคำตอบ 

```js

(getUserInput) =>  assert.equal( "Post แล้ว", getUserInput('url') );
   
```

# --solutions--

```html
// solution required
```
