---
id: 587d7fb1367417b2b2512bf3
title: Implement a Root-Level Request Logger Middleware
challengeType: 2
forumTopicId: 301514
dashedName: implement-a-root-level-request-logger-middleware
---

# --description--

ก่อนหน้านี้ คุณได้คุ้นเคยกับฟังก์ชันมิดเดิลแวร์ `express.static()` แล้ว ตอนนี้ได้เวลาดูว่ามิดเดิลแวร์คืออะไรในในเชิงลึกมากขึ้น ฟังก์ชันมิดเดิลแวร์คือฟังก์ชันที่รับ 3 อาร์กิวเมนต์ ได้แก่ ออบเจ็กต์คำขอ ออบเจ็กต์การตอบสนอง และฟังก์ชันถัดไปในวงจร request-response cycle ของแอปพลิเคชัน ซึ่งฟังก์ชันเหล่านี้รันโค้ดบางอย่างได้ แต่จะอาจมีไซด์เอฟเฟคหรือผลข้างเคียงในแอป และมักจะเพิ่มข้อมูลไปยังออบเจกต์คำขอหรืออ็อบเจ็กต์การตอบสนอง พวกเขายังสามารถหยุดการทำงานด้วยการส่งคำตอบเมื่อตรงตามเงื่อนไขบางประการได้อีกด้วย แต่ถ้า หากพวกเขาไม่ส่งการตอบกลับเมื่อทำเสร็จสิ้นแล้ว พวกเขาจะเริ่มดำเนินการสแต็กในฟังก์ชันถัดไป สิ่งเหล่านี้จะเรียกว่า อาร์กิวเมนต์ที่ 3 หรือ `next()`

คุณสามารถดูตัวอย่างได้ตามนี้:

```js
function(req, res, next) {
  console.log("I'm a middleware...");
  next();
}
```

สมมติว่า คุณติดตั้งฟังก์ชันบนเส้นทาง เมื่อคำขอตรงกับเส้นทาง มันจะแสดงผลสตริง “I'm a middleware...” หลังนั้นจึงเรียกใช้ฟังก์ชันถัดไปในสแต็ก ในแบบฝึกหัดนี้ คุณจะต้องสร้างมิดเดิลแวร์ระดับรูท (root-level) ดังที่คุณเห็นในโจทย์ Challengeข้อที่ 4 เพื่อที่จะติดตั้งฟังก์ชันมิดเดิลแวร์ที่ระดับรูท คุณสามารถใช้วิธี `app.use(<mware-function>)` ได้ สำหรับกรณีนี้ ฟังก์ชันจะทำงานตามคำขอทั้งหมด และคุณก็สามารถกำหนดเงื่อนไขเฉพาะเพิ่มเติมได้เช่นกัน ยกตัวอย่างเช่น ต้องการให้ฟังก์ชันทำงานเฉพาะสำหรับคำขอ POST คุณสามารถใช้ `app.post(<mware-function>)` ได้เลย นอกจากนี้ยังมีเวิร์บ http ที่ทำงานคล้ายกันทั้งหมด (GET, DELETE, PUT, …)

# --instructions--

ในการสร้าง ล็อคเกอร์แบบง่าย ควรเข้าสู่คอนโซลด้วยสตริงที่มีรูปแบบต่อไปนี้ `method path - ip` ในทุกคำขอ เช่น `GET /json - ::ffff:127.0.0.1` เราจะเห็นว่ามีช่องว่างระหว่าง `method` กับ `path` และเส้นประที่แยก `path` กับ `ip` ที่ล้อมรอบด้วยช่องว่างทั้งสองด้าน คุณสามารถใช้วิธีการขอ(request method, http method), เส้นทางที่เกี่ยวข้องกัน และ IP จากอ็อบเจกต์รีเควส ด้วยการใช้ `req.method`, `req.path` และ `req.ip` จำเอาไว้เสมอว่า ต้องเรียกใช้ `next()` ทุกครั้งเมื่อคุณทำเสร็จ ไม่อย่างนั้นเซิร์ฟเวอร์ของคุณจะติดหรือรันไม่ออก และอย่างสุดท้าย ต้องมั่นใจว่าเปิดล็อคแล้ว และต้องเช็คด้วยว่าเกิดอะไรขึ้นเมื่อมีคำขอบางอย่างมาถึง

**Note:** การประเมินฟังก์ชั่นนั้น จะทำตามลำดับที่ปรากฏในโค้ด ซึ่งรวมถึงกรณีของมิดเดิลแวร์ด้วย ถ้าหากคุณต้องการให้มันใช้งานได้กับทุกเส้นทาง ก็ควรติดตั้งไว้ในลำดับก่อนหน้านั้นด้วย

# --hints--

มิดเดิลแวร์ระดับรูทล็อคเกอร์ควรเปิดใช้งานอยู่

```js
(getUserInput) =>
  $.get(getUserInput('url') + '/_api/root-middleware-logger').then(
    (data) => {
      assert.isTrue(
        data.passed,
        'root-level logger is not working as expected'
      );
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

# --solutions--

```js
/**
  Backend challenges don't need solutions, 
  because they would need to be tested against a full working project. 
  Please check our contributing guidelines to learn more.
*/
```
