---
id: a3f503de51cf954ede28891d
title: Find the Symmetric Difference
challengeType: 5
forumTopicId: 301611
dashedName: find-the-symmetric-difference
---

# --description--

ในการดำเนินการทางคณิตศาตร์ จากแผนภาพของ เวนน์-ออยเลอร์ กำหนดให้ (`△` or `⊕`) คือ <dfn>symmetric difference</dfn> ระหว่าง 2 Sets
โดยเป็นการ สมาชิกของเซตมารวมกัน (Union) แต่ไม่นับเอาสมาชิกมีเหมือนกัน (Intersection) ตัวอย่างเช่น

`A = {1, 2, 3}` และ `B = {2, 3, 4}`, `A △ B = {1, 4}`

Note เพิ่มเติม. จากตัวอย่างด้านบน A ∪ B = {1, 2, 3, 4}  และ A ∩ B = {2, 3}
A △ B = (A ∪ B) - (A ∩ B)


Symmetric difference เป็นการดำเนินการแบบ binary กล่าวง่ายๆคือ สามารถดำเนินการได้ทีละคู่เท่านั้น
ดังนั้นการหาค่าที่มีตั้งแต่ 3 องค์ประกอบขึ้นไป  จะต้องดำเนินการทีละคู่ เช่น กำหนดให้ `A = {1, 2, 3}`, `B = {2, 3, 4}` และ  `C = {2, 3}`


`A △ B △ C` จะกลายเป็น (A △ B) △ C  และเนื่องจาก `A △ B = {1, 4}` และ  `C = {2, 3}`

จะได้เป็น `A △ B △ C    =   (A △ B) △ C     =    {1, 4} △ {2, 3}      =      {1, 2, 3, 4}`




# --instructions--

สร้างฟังก์ชันที่รับอาร์เรย์ตั้งแต่สองอาร์เรย์ขึ้นไป และส่งคืนอาร์เรย์ที่มีความแตกต่างของจำนวนสมาชิก โดยที่ อาร์เรย์ที่ส่งคืนต้องมี เฉพาะค่าที่ไม่ซ้ำกัน (*ค่าใน อาร์เรย์ ไม่ซ้ำกัน*).

# --hints--

`sym([1, 2, 3], [5, 2, 1, 4])` ควรได้ผลลัพธ์เป็น `[3, 4, 5]`.

```js
assert.sameMembers(sym([1, 2, 3], [5, 2, 1, 4]), [3, 4, 5]);
```

`sym([1, 2, 3], [5, 2, 1, 4])` ต้องมีสมาชิกเพียง 3 ตัว


```js
assert.equal(sym([1, 2, 3], [5, 2, 1, 4]).length, 3);
```

`sym([1, 2, 3, 3], [5, 2, 1, 4])` ควรได้ผลลัพธ์เป็น `[3, 4, 5]`.

```js
assert.sameMembers(sym([1, 2, 3, 3], [5, 2, 1, 4]), [3, 4, 5]);
```

`sym([1, 2, 3, 3], [5, 2, 1, 4])` ต้องมีสมาชิกเพียง 3 ตัว

```js
assert.equal(sym([1, 2, 3, 3], [5, 2, 1, 4]).length, 3);
```

`sym([1, 2, 3], [5, 2, 1, 4, 5])` ควรได้ผลลัพธ์เป็น `[3, 4, 5]`.

```js
assert.sameMembers(sym([1, 2, 3], [5, 2, 1, 4, 5]), [3, 4, 5]);
```

`sym([1, 2, 3], [5, 2, 1, 4, 5])` ต้องมีสมาชิกเพียง 3 ตัว

```js
assert.equal(sym([1, 2, 3], [5, 2, 1, 4, 5]).length, 3);
```

`sym([1, 2, 5], [2, 3, 5], [3, 4, 5])` ควรได้ผลลัพธ์เป็น `[1, 4, 5]`

```js
assert.sameMembers(sym([1, 2, 5], [2, 3, 5], [3, 4, 5]), [1, 4, 5]);
```

`sym([1, 2, 5], [2, 3, 5], [3, 4, 5])` ต้องมีสมาชิกเพียง 3 ตัว

```js
assert.equal(sym([1, 2, 5], [2, 3, 5], [3, 4, 5]).length, 3);
```

`sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5])` ควรได้ผลลัพธ์เป็น `[1, 4, 5]`.

```js
assert.sameMembers(sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5]), [1, 4, 5]);
```

`sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5])` ต้องมีสมาชิกเพียง 3 ตัว

```js
assert.equal(sym([1, 1, 2, 5], [2, 2, 3, 5], [3, 4, 5, 5]).length, 3);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3])` ควรได้ผลลัพธ์เป็น `[2, 3, 4, 6, 7]`.

```js
assert.sameMembers(
  sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3]),
  [2, 3, 4, 6, 7]
);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3])` ต้องมีสมาชิกเพียง 5 ตัว

```js
assert.equal(
  sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3]).length,
  5
);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3], [5, 3, 9, 8], [1])` ควรได้ผลลัพธ์เป็น `[1, 2, 4, 5, 6, 7, 8, 9]`.

```js
assert.sameMembers(
  sym(
    [3, 3, 3, 2, 5],
    [2, 1, 5, 7],
    [3, 4, 6, 6],
    [1, 2, 3],
    [5, 3, 9, 8],
    [1]
  ),
  [1, 2, 4, 5, 6, 7, 8, 9]
);
```

`sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3], [5, 3, 9, 8], [1])` ต้องมีสมาชิกเพียง 8 ตัว

```js
assert.equal(
  sym([3, 3, 3, 2, 5], [2, 1, 5, 7], [3, 4, 6, 6], [1, 2, 3], [5, 3, 9, 8], [1])
    .length,
  8
);
```

# --seed--

## --seed-contents--

```js
function sym(args) {
  return args;
}

sym([1, 2, 3], [5, 2, 1, 4]);
```

# --solutions--

```js
function sym() {
  var arrays = [].slice.call(arguments);
  return arrays.reduce(function (symDiff, arr) {
    return symDiff.concat(arr).filter(function (val, idx, theArr) {
      return theArr.indexOf(val) === idx
        && (symDiff.indexOf(val) === -1 || arr.indexOf(val) === -1);
    });
  });
}
sym([1, 2, 3], [5, 2, 1, 4]);
```
