---
id: 587d7b85367417b2b2512b38
title: Catch Use of Assignment Operator Instead of Equality Operator
challengeType: 1
forumTopicId: 301191
dashedName: catch-use-of-assignment-operator-instead-of-equality-operator
---

# --description--

โปรแกรมที่มีการทำงานแยกเป็นเงื่อนไข เช่น โปรแกรมที่จำทำงานอย่างนึงหากเข้าเงื่อนไขนึง และทำงานอีกแบบถ้าไม่เข้าเงื่อนไข  
โปรแกรมที่มีการทำงานแยกเป็นเงื่อนไขนี้มักมีการใช้คำสั่ง `if`, `else if` และ `else` ซึ่งเงื่อนไขบางครั้งก็เป็นการตรวจสอบว่าค่าสองตัวเท่ากันหรือไม่

ถ้าเราจะเปลี่ยน logic นี้เป็นคำพูดก็จะได้ประมาณว่า "ถ้า x แล้ว y หลังจากนั้น ..." ซึ่งถ้าเราเขียนโค้ดตามคำพูดแบบนี้ แล้วใช้เครื่องหมายเท่ากับ `=` หรือ assignment operator ตรงตามตัวเลยจะทำให้โค้ดเราทำงานไม่ถูกต้อง

จากที่เราได้เรียนมาก่อนหน้านี้ assignment operator (`=`) ใน JavaScript ใช้กำหนดค่าให้กับตัวแปร ส่วน `==` และ `===` ใช้ในการตรวจสอบเงื่อนไขว่าเท่ากันหรือไม่ (`===` ใช้ในการตรวจสอบ strict equality หมายถึงต้องเหมือนกันทั้งค่าและชนิดข้อมูล)

โค้ดด้านล่างนี้กำหนดให้ตัวแปร `x` มีค่าเป็น 2 ซึ่งถือว่าเป็นค่า `true` แต่ค่าเกือบทุกค่าใน JavaScript ก็ถือเป็น `true` ทั้งนั้น 
ยกเว้นค่าที่เรียกว่า "falsy" value คือ `false`, `0`, `""` (string ว่าง), `NaN`, `undefined` และ `null` ที่จะถือว่าเป็น false

```js
let x = 1;
let y = 2;
if (x = y) {

} else {

}
```

จากตัวอย่างข้างบน โค้ดที่อยู่ในคำสั่ง `if` จะทั้งานเสมอไม่ว่าค่า `y` จะเป็นอะไรก็ตาม ยกเว้นว่าถ้า `y` เป็น falsy value ดังนั้น โค้ดภายใน `else` ที่เราคาดหวังให้ทำงานก็จะไม่ถูกทำงาน

สาเหตุที่เป็นแบบนั้นเพราะใน `()` ของ `if` นั้นจะทำการกำหนดค่าของ `y` ให้กับ `x` เมื่อกำหนดค่าแล้วโปรแกรมจะเห็นคำสั่ง `if` เป็นรูปแบบนี้แทน `if (x)` ซึ่งจริงๆแล้วเราอาจจะอยากให้มันเป็น `if (x==y)`

# --instructions--

จงปรับแก้เงื่อนไข เพื่อให้โปรแกรมทำงานได้ถูกต้อง และ `result` แสดงผลลัพธ์ได้ถูกต้อง

# --hints--

ต้องแก้เงื่อนไขเพื่อให้ตรวจสอบว่าค่า `x` และ `y` เท่ากันหรือไม่ แทนที่จะเป็นการกำหนดค่าให้ตัวแปร

```js
assert(result == 'Not equal!');
```

เงื่อนไขควรใช้ `==` หรือ `===` เพื่อตรวจสอบว่าค่า `x` และ `y` เท่ากันหรือไม่

```js
assert(code.match(/x\s*?===?\s*?y/g));
```

# --seed--

## --seed-contents--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x = y) {
  result = "Equal!";
} else {
  result = "Not equal!";
}

console.log(result);
```

# --solutions--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x === y) {
 result = "Equal!";
} else {
 result = "Not equal!";
}

console.log(result);
```
