---
id: cf1111c1c11feddfaeb3bdef
title: Add Two Numbers with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` เป็นประเภทข้อมูลใน JavaScript ซึ่งแทนข้อมูลตัวเลข

ขั้นตอนนี้ลองนำตัวเลข 2 ตัวมาบวกกันโดยใช้ JavaScript

JavaScript ใช้สัญลักษณ์ `+` วางไว้ระหว่างตัวเลขสองตัว เพื่อทำการบวก


**Example:**

```js
myVar = 5 + 10;
```

`myVar` ตอนนี้มีค่า `15`

# --instructions--

จงเปลี่ยน `0` เพื่อให้ผลรวมเท่ากับ `20`

# --hints--

`sum` ควรเท่ากับ `20`

```js
assert(sum === 20);
```

คุณควรใช้ `+` operator

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
var sum = 10 + 0;
```

# --solutions--

```js
var sum = 10 + 10;
```
