---
id: cf1111c1c11feddfaeb3bdef
title: Add Two Numbers with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` เป็นประเภทข้อมูลใน JavaScript ซึ่งหมายถึงข้อมูลที่เป็นตัวเลข

ในบทเรียนนี้ให้ลองนำตัวเลข 2 ตัวมาบวกกันโดยใช้ JavaScript

การบวกใน JavaScript ทำได้โดยการวางเครื่องหมาย `+` ไว้ระหว่างตัวเลขสองตัว


**Example:**

```js
myVar = 5 + 10;
```

ตอนนี้ตัวแปร `myVar` มีค่าเท่ากับ `15`

# --instructions--

จงเปลี่ยน `0` เพื่อให้ผลรวมเท่ากับ `20`

# --hints--

ตัวแปร `sum` ควรมีค่าเท่ากับ `20`

```js
assert(sum === 20);
```

คุณควรใช้เครื่องหมาย `+`

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
var sum = 10 + 0;
```

# --solutions--

```js
var sum = 10 + 10;
```
