---
id: cf1111c1c11feddfaeb6bdef
title: Divide One Number by Another with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cqkbdAr'
forumTopicId: 17566
dashedName: divide-one-number-by-another-with-javascript
---

# --description--

เราสามารถหารตัวเลขหนึ่งด้วยเลขอีกตัวหนึ่งได้เช่นกัน

JavaScript ใช้เครื่องหมาย `/` สำหรับการหาร

**Example**

```js
myVar = 16 / 2;
```

ตัวแปร `myVar` ตอนนี้จะมีค่าเท่ากับ `8`

# --instructions--

จงเปลี่ยน `0` เพื่อให้ตัวแปร `quotient` มีค่าเท่ากับ `2`

# --hints--

ตัวแปร `quotient` ควรมีค่าเท่ากับ 2

```js
assert(quotient === 2);
```

คุณควรใช้เครื่องหมาย `/`


```js
assert(/\d+\s*\/\s*\d+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'quotient = '+z;})(quotient);
```

## --seed-contents--

```js
var quotient = 66 / 0;
```

# --solutions--

```js
var quotient = 66 / 33;
```
