---
id: bd7993c9c69feddfaeb7bdef
title: Multiply Two Decimals with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/ce2GeHq'
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

ใน JavaScript คุณสามารถคำนวณเลขทศนิยมได้เช่นเดียวกับเลขจำนวนเต็ม

ลองคูณเลขทศนิยมสองตัวเพื่อหาผลคูณ


# --instructions--

จงเปลี่ยน `0.0` เพื่อให้ได้ผลลัพย์เท่ากับ `5.0`

# --hints--

ตัวแปร `product` ควรเท่ากับ `5.0`

```js
assert(product === 5.0);
``` 

คุณควรใช้ `*` operator

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(y){return 'product = '+y;})(product);
```

## --seed-contents--

```js
var product = 2.0 * 0.0;
```

# --solutions--

```js
var product = 2.0 * 2.5;
```
