---
id: cf1231c1c11feddfaeb5bdef
title: Multiply Two Numbers with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3y3Aq'
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

เราสามารถคูณตัวเลขหนึ่งกับอีกเลขหนึ่งได้เช่นกัน

JavaScript ใช้เครื่องหมาย `*` สำหรับการคูณ

**Example**

```js
myVar = 13 * 13;
```

ตัวแปร `myVar` จะมีค่าเท่ากับ `169`.

# --instructions--

จงเปลี่ยน `0` เพื่อให้ผลคูณมีค่าเท่ากับ `80`.

# --hints--

ตัวแปร `product` ควรมีค่าเท่ากับ 80

```js
assert(product === 80);
```

คุณควรใช้เครื่องหมาย `*`

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'product = '+z;})(product);
```

## --seed-contents--

```js
var product = 8 * 0;
```

# --solutions--

```js
var product = 8 * 10;
```
