---
id: cf1111c1c11feddfaeb4bdef
title: Subtract One Number from Another with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3yQtk'
forumTopicId: 18314
dashedName: subtract-one-number-from-another-with-javascript
---

# --description--

เราสามารถลบตัวเลขหนึ่งออกจากอีกจำนวนหนึ่งได้

JavaScript ใช้สัญลักษณ์ `-` สำหรับการลบ

**Example**

```js
myVar = 12 - 6;
```

`myVar` ควรมีค่า `6`.

# --instructions--

จงเปลี่ยน `0` เพื่อ `difference` มีค่าเป็น `12`

# --hints--

ตัวแปร `difference` ควรเท่ากับ 12.

```js
assert(difference === 12);
```

คุณควรนำตัวเลขหนึ่งตัวมาลบออกจาก 45

```js
assert(/difference=45-33;?/.test(__helpers.removeWhiteSpace(code)));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'difference = '+z;})(difference);
```

## --seed-contents--

```js
var difference = 45 - 0;
```

# --solutions--

```js
var difference = 45 - 33;
```
